# Week 8 Practice: Regular Expressions

```python
import re
```

### Problem 1
You are given a string consisting only of digits, commas, and periods.

Write a function that uses `re.split` to split the string at its commas and periods, returning a list of the number's numeric parts.

It’s guaranteed that every comma and every period is preceeded and followed by a digit.

<!-- grade_id: cell-84eb6b2a9fd47921, points: 0, grade: false, solution: true -->
```python
def numeric_parts(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-408521412c7d67d8, points: 1, grade: true, solution: false -->
```python
assert numeric_parts("100,000,000.000") == ["100", "000", "000", "000"]
```

### Problem 2

You are given a string .
Your task is to verify that  is a floating point number.

In this task, a valid float number must satisfy all of the following requirements:
- can start with `+`, `-` or `.` symbol.
- must contain at least one decimal value.
- must have exactly one `.` symbol.

<!-- grade_id: cell-453e7b882a366588, points: 0, grade: false, solution: true -->
```python
def is_float(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ca3384b22f50663d, points: 1, grade: true, solution: false -->
```python
assert not is_float("4.0O0")
assert is_float("-1.00")
assert is_float ("+4.54")
assert not is_float("SomeRandomStuff")
assert is_float("12.0")
assert not is_float("12.")
assert is_float(".7")
assert is_float("-0.7")
```