###### Week 6 Numpy Practice Problems

```python
%load_ext pycodestyle_magic
%flake8_on --ignore E266 --max_line_length 99

import numpy as np
```

### Problem 1
Convert an array of Celsius measurements to Fahrenheit.

<!-- grade_id: cell-91b2c99aa4c4ff4c, points: 0, grade: false, solution: true -->
```python
def convert_to_fahrenheit(centigrade_temps):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5865bd78c4dcf6aa, points: 1, grade: true, solution: false -->
```python
assert np.all(convert_to_fahrenheit(np.abs(np.zeros(4)) - 32.) < 1e-8)
assert np.all(convert_to_fahrenheit(np.abs(np.ones((10, 10)) * 100) - 212.) < 1e-8)
```

### Problem 2
Use numpy to create a function that adds a border filled with zeros around an existing array.

*Hint*: Use numpy's `pad` function. See https://numpy.org/doc/stable/reference/generated/numpy.pad.html

<!-- grade_id: cell-2466342b05a17106, points: 0, grade: false, solution: true -->
```python
def add_border(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fb9bae34e3908aed, points: 1, grade: true, solution: false -->
```python
x = np.ones((1, 1))
claimed = add_border(x)
correct = np.zeros((3, 3))
correct[1, 1] = 1.0
assert claimed.shape == correct.shape
assert np.all(np.abs(claimed - correct) < 1e-8)
```

### Problem 3
Create an $n \times n$ numpy array filled with a checkerboard pattern of zeros and ones. The first entry should be the integer `0`. Assume $n \ge 2$.

<!-- grade_id: cell-9982d08b4ea7057d, points: 0, grade: false, solution: true -->
```python
def make_checkerboard(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e6f9eb1adf863ce2, points: 1, grade: true, solution: false -->
```python
correct = \
    [[0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0]]
assert np.all(correct == make_checkerboard(8))
```