# Week 10 Practice

```python
%load_ext pycodestyle_magic
%flake8_on --ignore E266 --max_line_length 99
```

### Recap

Last class we implmented functions called `f_torch` that evaluated $f(x) = 3x^\intercal x - x_1 - 4$.

```python
import torch
import numpy as np


def f_torch(x):
    return 3 * x.dot(x) - x[0] - 4
```

### Problem 1: Minimizing our function with gradient descent
Write a function that minimizes $f$ using gradient descent, by calling by calling `f_torch`. The function should take two arguments: an initial iterate `x0` and a number of gradient descent iterations `num_iters`. It should return the minimizer.

<!-- grade_id: cell-e0aba96eeaf8460e, points: 0, grade: false, solution: true -->
```python
def minimize(x0, num_iters, step_size=0.1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3453f58f7aaffae2, points: 1, grade: true, solution: false -->
```python
x_opt = minimize(torch.zeros(3), 30)
assert (x_opt - torch.Tensor([0.1667, 0.0000, 0.0000])).norm() < 1e-3
```

### Problem 2: Working with PyTorch datasets
Find the mean brightness of CIFAR-10 images by iterating through the dataset and summing the pixel intensities.

```python
from torchvision import datasets
import torchvision.transforms as transforms


trainset = datasets.CIFAR10(root='./data',
                            train=True,
                            download=True,
                            transform=transforms.ToTensor())
```

```python
from matplotlib import pyplot as plt


def imshow(img):
    npimg = img.numpy()
    plt.imshow(np.transpose(npimg, (1, 2, 0)))
    plt.show()


for img, label in trainset:
    imshow(img)
    break
```

<!-- grade_id: cell-9f3da1d5dcc25501, points: 0, grade: false, solution: true -->
```python
def mean_brightness_directly(dataset):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-321bdf2544b88758, points: 1, grade: true, solution: false -->
```python
assert abs(mean_brightness_directly(trainset) - 0.4733) < 1e-4
```

### Problem 3: Data loaders
Find the mean brightness of CIFAR-10 images by iterating through the dataset using a DataLoader with a batch size of 256.

<!-- grade_id: cell-36a333d71d7c1207, points: 0, grade: false, solution: true -->
```python
from torch.utils.data import DataLoader


def mean_brightness_dataloader(dataset):
    loader = DataLoader(trainset, batch_size=256, shuffle=True)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-82c31d17f7ec115a, points: 1, grade: true, solution: false -->
```python
assert abs(mean_brightness_dataloader(trainset) - 0.4733) < 1e-4
```

### Problem 4: Minimizing a loss function using SGD
Find the mean brightness of CIFAR-10 images in two ways: by iterating through the dataset and summing the pixel intensities and with stochastic gradient descent, by minimizing squared error.

<!-- grade_id: cell-8f04dd4d34016148, points: 0, grade: false, solution: true -->
```python
def mean_brightness_sgd(dataset):
    loader = DataLoader(trainset, batch_size=256, shuffle=True)
    cur_mean = torch.zeros(1)
    cur_mean.requires_grad_(True)
    # YOUR CODE HERE
    raise NotImplementedError()

    return cur_mean.data.item()
```

<!-- grade_id: cell-c70dae2beab0f068, points: 1, grade: true, solution: false -->
```python
assert abs(mean_brightness_sgd(trainset) - 0.4733) < 1e-2
```