### Lecture 3 Practice: Dictionaries

### Question 1
Define a function called `dict_square`, which takes a positive integer $n$ as its only argument, and returns a dictionary that maps each number from 1 to $n$ to its square.

<!-- grade_id: cell-f08d8946eba9b834, points: 0, grade: false, solution: true -->
```python
def dict_square(n):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-03e3f94708532e9c, points: 1, grade: true, solution: false -->
```python
assert(dict_square(1) == {1: 1})
assert(dict_square(4) == {1: 1, 2: 4, 3: 9, 4: 16})
```

### Question 2
Define a function called `switch_key_value`, which takes a dictionary as its only argument. Switch the placement of key and value of the dictionary and return it as a new dictionary.

<!-- grade_id: cell-6c67d579c481de9b, points: 0, grade: false, solution: true -->
```python
def switch_key_value(d):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a0b5ecba31ff8afb, points: 1, grade: true, solution: false -->
```python
assert(switch_key_value({}) == {})
assert(switch_key_value({'a': 1, 'b': 2}) == {1: 'a', 2: 'b'})
```

### Question 3
Define a function called `sort_list_dict`, which takes a dictionary whose values are lists as its only argument. Sort all these lists alphabetically return them in a new dictionary (with the same keys).

<!-- grade_id: cell-996c61e79b68ca6d, points: 0, grade: false, solution: true -->
```python
def sort_list_dict(d):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f752b519fbf2ac78, points: 1, grade: true, solution: false -->
```python
assert(sort_list_dict({}) == {})

assert(sort_list_dict({'a': []}) == {'a': []})

input_dict = {'a': [4, 1, 2, 3], 'b': [3, 2, 1]}
desired_output = {'a': [1, 2, 3, 4], 'b': [1, 2, 3]}
assert(sort_list_dict(input_dict) == desired_output)

input_dict = {'S': [2, 3, 1], 'T': [5, 1, 2], 'D': [3, 2, 4]}
desired_output = {'S': [1, 2, 3], 'T': [1, 2, 5], 'D': [2, 3, 4]}
assert(sort_list_dict(input_dict) == desired_output)
```

### Question 4

Define a function called combine, which takes two dictionaries as its argument. Combine the two dictionaries by adding values for common keys and return results as a new dictionary. Your function should raise an appropriate error in the event that the input is not dictionary.

<!-- grade_id: cell-8fab97e76663cc9e, points: 0, grade: false, solution: true -->
```python
def combine(d1, d2):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-39e252004f8545a6, points: 1, grade: true, solution: false -->
```python
assert(combine({}, {}) == {})

assert(combine({'a': 100}, {'b': 100}) == {'a': 100, 'b': 100})

input_dicts = ({'a': 100, 'b': 100, 'c': 300}, {'b': 100})
assert(combine(*input_dicts) == {'a': 100, 'b': 200, 'c': 300})

input_dicts = ({'a': 100, 'b': 200}, {'a': 100, 'c': 200})
assert(combine(*input_dicts) == {'a': 200, 'b': 200, 'c': 200})

try:
    combine([1], {})
    assert(False)
except TypeError:
    assert(True)
except Exception:
    assert(False)
```