# Homework 5 (28 points)

name: XX

email: XX@umich.edu

This homework assignment took me XX hours in total to complete. (Please help us to gauge the difficulty of the assignment.)

## Collaboration Disclosure

In the cell below, please list *everyone* with whom you discussed any of the homework problems, excluding only the GSIs and the course instructor. 

If you did not discuss the homework with anyone else, write __"I did not discuss this homework with anyone."__

Even if you discuss questions with other, the code you submit must be only yours. All work is checked with the [MOSS plagiarism detector](https://theory.stanford.edu/~aiken/moss/).

I discussed XX with XX. 

## Submission Instructions
Your homework solutions should be written entirely in this Jupyter notebook file. Once it contains your solutions, you should submit this notebook through Canvas. Before submitting, please make sure make sure __Cells->Run__ executes without errors; errors in your code cause problems for our auto-grader. Thus, even though we may ask you to check for errors in your functions, you should not include in your submission any examples of your functions actually raising those errors.

## Nbgrader

We will be using `nbgrader` to grade your jupyter notebook. Additional tests will be used at grading time, you are encouraged to read the problem carefully and verify your code covers all appropriate cases.

**Be careful:** If a jupyter notebook cell takes longer than `60s` to run, the autograder will not grade it and you will receive zero credit for that question.

## Error checking

You do not need to do error checking (raising errors, etc.) in your code unless we explicitly ask you to so in a problem.

## Coding style
Please follow the [PEP8 guidelines](https://www.python.org/dev/peps/pep-0008/). The `pycodestyle_magic` package will remind you if you are not. To install it, run
```
pip install flake8 pycodestyle_magic
```
Once it is installed the following two Jupyter directives will enable syntax checking:

```python
%load_ext pycodestyle_magic
%flake8_on --ignore E266 --max_line_length 99
```

```python
import numpy as np
import itertools
import functools
import random
import math
import scipy.special
```

### Problem 1: Basic Review Questions (9 points)

1. <b>(2 points)</b> Write a function called `collatz` takes a positive integer `num` as its argument, and returns a list. Your function should perform appropriate error checking for it. The list is generated as follows: The first element of the list is `num`. Continue doing the following things: if the last element of the list, `number`, is even, then the next element of the list is `number` / 2; if the last element of the list, `number`, is odd, then the next element of the list is 3 * number + 1. Stop when the last element of the list is 1. (Note: We guarantee the length of the list is finite.)

<!-- grade_id: cell-108419d75c8da43a, points: 0, grade: false, solution: true -->
```python
def collatz(num):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: collatz, points: 0.5, grade: true, solution: false -->
```python
assert collatz(3) == [3, 10, 5, 16, 8, 4, 2, 1]
```

<!-- grade_id: collatz-value-error, points: 0.5, grade: true, solution: false -->
```python
assert type(collatz(5)) == list
```

<!-- grade_id: collatz-type-error, points: 0.5, grade: true, solution: false -->
```python
try:
    collatz(-1)

except ValueError:
    pass

else:
    raise AssertionError("Function allows for negative int as input but it shouldn't.")
```

<!-- grade_id: cell-d6cb9403c8069e69, points: 0.5, grade: true, solution: false -->
```python
assert collatz(2) == [2, 1]
```

2. <b>(1.5 points)</b> __Casino Game:__ This game in the casino consists of tossing a die with N faces twice. There is a number written on each face of the die (these numbers are not necessarily distinct). In order to win, you must get the number $A$ on the first toss and the number $B$ on the second toss of the die. The excited viewers want to know the probability that you will win the game. Define a function `casino_game` that takes a list of integers  `n`, an integer `a` and an integer `b` as its arguments and returns the probability that you win the game. Your function should perform appropriate error checking. Assume that you get each face of the die with the same probability on each toss and that tosses are mutually independent.

<!-- grade_id: cell-611c0aafe5caa0cc, points: 0, grade: false, solution: true -->
```python
def casino_game(n, a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: casino-game-easy, points: 0.5, grade: true, solution: false -->
```python
assert casino_game([1, 2, 3, 4], 1, 1) == 0.0625
```

<!-- grade_id: cell-42f199155837eb9e, points: 0.5, grade: true, solution: false -->
```python
assert casino_game([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0], 0, 0) == 1.0
```

<!-- grade_id: casino-game-type-errors, points: 0.5, grade: true, solution: false -->
```python
try:
    casino_game(0, 1, 1)

except TypeError:
    pass

else:
    raise AssertionError("Function allows for int input where a list is expected.")
```

3. <b>(2 points)</b> Write a function `dict_comp` that takes a positive integer `n` as its argument and returns a dictionary that contains the numbers (between 1 and `n`) as its keys and the first `n` composite numbers as its values. Your function should perform appropriate error checking for it. Recall that a composite number is any integer $p > 1$ which is not prime; therefore, 4 is the first composite number.

<!-- grade_id: cell-5129b62a1c02655b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: dict-comp-easy, points: 0.5, grade: true, solution: false -->
```python
assert dict_comp(5) == {1: 4, 2: 6, 3: 8, 4: 9, 5: 10}
```

<!-- grade_id: dict-comp-hard, points: 0.5, grade: true, solution: false -->
```python
assert type(dict_comp(3)) == dict
```

<!-- grade_id: dict-type-errors, points: 0.5, grade: true, solution: false -->
```python
try:
    dict_comp([0])

except TypeError:
    pass

else:
    raise AssertionError("Function allows for string input where an int is expected.")
```

<!-- grade_id: dict-value-errors, points: 0.5, grade: true, solution: false -->
```python
try:
    dict_comp(0)

except ValueError:
    pass

else:
    raise AssertionError("Value allows for input that is not a positive integer.")
```

4. <b>(1 point)</b> Write a one-line expression that computes the sum of the first 10 even square numbers (starting from $2^2$). For ease of grading, please assign the output of this expression to a variable called `even_squares`. Please make sure that the cell that defines even_squares is the one directly above the test case for this question.

<!-- grade_id: even-squares, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6f5cd818d00af6fd, points: 1, grade: true, solution: false -->
```python
# Make sure that the cell that defines even_squares is directly above this one.
assert type(even_squares) == int or type(even_squares) == float
```

5. <b>(1 point)</b> Write a generator expression that enumerates the tetrahedral numbers. The n-th tetrahedral number ($n = 1,\ 2,\ \cdots$) is given by $Tn = \binom{n+2}{3}$， where $\binom{x}{y}$ is the binomial coefficient
$$\binom{x}{y} = \frac{x!}{y!(x-y)!}$$
For ease of grading, please assign this generator expression to a variable called `tetra`. Hint: you may find it useful to define a generator that enumerates the positive integers.

<!-- grade_id: cell-2ade7aa604144668, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: tetra, points: 1, grade: true, solution: false -->
```python
import types
assert isinstance(tetra, types.GeneratorType)
```

6. <b>(1.5 points)</b> You are given a string `s`, which contains alphanumeric characters only (this means you can assume that the string will have no spaces nor punctuation). Your task is to sort the string `s` in the following manner:  
 - All sorted lowercase letters are ahead of uppercase letters.
 - All sorted uppercase letters are ahead of digits.
 - All sorted odd digits are ahead of sorted even digits.

Write a function called `sort_str`, which takes the string `s` as its argument, and return the result. Your function should perform appropriate error checking for the input.

<!-- grade_id: cell-752e6972babe6b23, points: 0, grade: false, solution: true -->
```python
def sort_str(s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: sort-str, points: 1, grade: true, solution: false -->
```python
assert sort_str('Sorting1234') == 'ginortS1324'
```

<!-- grade_id: sort-str-type-error, points: 0.5, grade: true, solution: false -->
```python
try:
    sort_str(123,)

except TypeError:
    pass

else:
    raise AssertionError("Function allows for int input where string inputs is expected.")
```

### Problem 2: Fantasy Game (5 points)
1. <b>(2 points)</b> __Fantasy Game Inventory:__ You are creating a fantasy video game. The data structure to model the player’s inventory will be a dictionary where the keys are string values describing the item in the inventory and the value is an integer value detailing how many of that item the player has. For example, the dictionary value `{'rope': 1, 'torch': 6, 'gold coin': 42, 'dagger': 1, 'arrow': 12}` means the player has 1 rope, 6 torches, 42 gold coins, and so on. Write a function named `displayInventory` that takes a dictionary `inventory` as its argument and prints out the inventory line by line. Using the example dictionary, it should print it out EXACTLY like the following:

Inventory: <br>
12 arrow <br>
42 gold coin <br>
1 rope <br>
6 torch <br>
1 dagger <br>
Total number of items: 62 <br>

__Note:__ Your function should raise TypeError if the input is not a dictionary.

__Note:__ You will not be graded on the order of the inventory items. The only print order that matters is that "Inventory:" is printed before all of the items and the total is printed after all of the items. Everything in between can occur in any order so long as it is accurate.

<!-- grade_id: cell-d6b22abac7bcf9ea, points: 0, grade: false, solution: true -->
```python
def display_inventory(inventory):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: display-inventory, points: 1, grade: true, solution: false -->
```python
inventory = {'rope': 1, 'torch': 6, 'gold coin': 42, 'dagger': 1, 'arrow': 12, 'map fragments': 3}
assert display_inventory(inventory) is None
```

<!-- grade_id: display-error, points: 1, grade: true, solution: false -->
```python
try:
    display_inventory(["1: rope", "2: torch"])

except TypeError:
    pass

else:
    raise AssertionError("Input of list doesn't raise proper error.")
```

2. <b>(3 points)</b> __List to Dictionary Function for Fantasy Game Inventory:__ Imagine that a vanquished dragon’s loot is represented as a list of strings like this: `dragonLoot = ['gold coin', 'dagger', 'gold coin', 'gold coin', 'ruby']`. Write a function named `addToInventory(inventory, added_items)`, where the `inventory` argument is a dictionary representing the player’s inventory (like in the previous question) and the `added_items` argument is a list like dragonLoot. Your function should check to make sure that the types of the input are appropriate (e.g., that inventory is a dictionary and addedItems should be a list of strings). The `addToInventory` function should return a dictionary that represents the updated inventory. __Note:__ the addedItems list can contain multiples of the same item.

<!-- grade_id: cell-d0f5b2d73004c4e5, points: 0, grade: false, solution: true -->
```python
def add_to_inventory(inventory, added_items):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: add-to-inv-easy, points: 1, grade: true, solution: false -->
```python
inv = {'gold coin': 42, 'rope': 1}
dragon_loot = ['gold coin', 'dagger', 'ruby']
assert add_to_inventory(inv, dragon_loot) == {'dagger': 1, 'gold coin': 43, 'rope': 1, 'ruby': 1}
```

<!-- grade_id: add-to-inv-hard, points: 1, grade: true, solution: false -->
```python
inv = {'gold coin': 42, 'rope': 1}
dragon_loot = ['gold coin', 'dagger', 'gold coin', 'gold coin', 'ruby']
correct_dict = {'dagger': 1, 'gold coin': 45, 'rope': 1, 'ruby': 1}
assert add_to_inventory(inv, dragon_loot) == correct_dict
```

<!-- grade_id: add-to-inv-errors, points: 0.5, grade: true, solution: false -->
```python
try:
    add_to_inventory([1, 1], ['a', 'b'])

except TypeError:
    pass

else:
    raise AssertionError("First input is not a dict.")
```

<!-- grade_id: add-to-inv-errors2, points: 0.5, grade: true, solution: false -->
```python
try:
    add_to_inventory({'a': 1, 'b': 1}, [2, 3])

except ValueError:
    pass
except TypeError:
    pass

else:
    raise AssertionError("Second input is a list of ints where a list of strings is expected.")
```

### Problem 3: More Fun with Strings (5 points)
1. <b>(2 points)</b> __Merge the Tools:__ Consider the following: A string, $s$, of length $n$  where $s=c_0c_1 \ldots c_{n-1}$. An integer, $k$ , where $k$ is a factor of $n$. We can split $s$ into $n/k$ subsegments where each subsegment, $t_i$ , consists of a contiguous block of $k$ characters in $s$. Then, use each $t_i$ to create string $u_i$ such that:  
 - The characters in $u_i$ are a subsequence of the characters in $t_i$.  
 - Any repeat occurrence of a character is removed from the string such that each character in $u_i$ occurs exactly once. In other words, if the character at some index $j$  in $t_i$  occurs at a previous index $< j$ in $t_i$, then do not include the character in string $u_i$.
  
 Write a function called `str_seg`, which takes the string `s` and the value `k` as its two argument, and return a list containing $\frac{n}{k}$ elements, where $i$-th element denotes the string $u_i$. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking. This means you can assume s and k will be of correct types in all the test cases and k will be a positive integer.

<!-- grade_id: cell-ad502ca49fb8f47e, points: 0, grade: false, solution: true -->
```python
def str_seg(s, k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: empty-string, points: 0.5, grade: true, solution: false -->
```python
assert type(str_seg("This is the first test case", 16)) == list
```

<!-- grade_id: str-seg, points: 1, grade: true, solution: false -->
```python
assert str_seg('AABCAAADA', 3) == ['AB', 'CA', 'AD']
```

<!-- grade_id: str-seg-edge-cases, points: 0.5, grade: true, solution: false -->
```python
assert str_seg("01234567890", 1) == ['0', '1', '2', '3', '4', '5', '6', '7', '8', '9', '0']
```

2. <b>(3 points)</b> __Simple Arithmetics:__ You are given multiple sequences of button presses of a simple pocket calculator that consist of digits and arithmetic operators. For each such sequence find the number it would produce on a pocket calculator's display. __Note that the pocket calculator evaluates the operators in the order in which they are given (i.e., there is no operator precedence (PEMDAS doesn't apply here)).__ Assume that the display of the calculator is large enough to show the result, and that its memory is sufficient to store all intermediate results. Write a function called `simple_arith`, which takes a string as its only argument, and returns the result as a float. In this problem, A string represents one sequence of button presses for a pocket calculator. The sequence consists of non-negative integers and arithmetic operators and ends with an equal sign. __It can also contain spaces to improve readability.__ __The operators only contain "+-*/"__, and the operator / represents integer division, rounded down. You may assume that no test case contains division by zero. __Hint:__ you may use `eval()` function. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking. This means you can assume that the input will always be a string with a valid mathematical expression in it.

<!-- grade_id: cell-c83de32d388931f8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: correct-return-type, points: 0.5, grade: true, solution: false -->
```python
assert type(simple_arith("5*5=")) == float
```

<!-- grade_id: simple-arith-simple, points: 0.5, grade: true, solution: false -->
```python
assert simple_arith('1+2+3 + 3 + 1*2 =') == 20.0
```

<!-- grade_id: simple-arith-digits, points: 0.5, grade: true, solution: false -->
```python
assert simple_arith('100-99-10+5=') == -4.0
```

<!-- grade_id: equals-only, points: 0.5, grade: true, solution: false -->
```python
assert simple_arith('9-9/5-5=') == -5.0
```

<!-- grade_id: simple-arith-hard, points: 1, grade: true, solution: false -->
```python
assert simple_arith("5*5/4 =") == 6.0
```

### Problem 4: More Practice (9 points)

1. <b>(2 points)</b> You are given a string $s$. Suppose a character '$c$' occurs consecutively $x$ times in the string. __In this problem, we assume each character is in '0123456789', representing a number between 0 and 9.__ Replace these consecutive occurrences of the character '$c$' with $(x,c)$ in the string and store into a list. Write a function called `comp_str`, which takes the string `s` as its argument, and return the list containing the replacements. For example, the result of `comp_str('1222311')` is `[(1, 1), (3, 2), (1, 3), (2, 1)]`. Reason: First, the character '1' occurs only once, so it is replaced by (1,1). Then the character '2' occurs three times, and it is replaced by (3,2) and so on. __Hint:__ _itertools.groupby_ might be helpful in this question. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking.

<!-- grade_id: cell-988506a92b718ef2, points: 0, grade: false, solution: true -->
```python
def comp_str(s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: comp-str-return, points: 0.5, grade: true, solution: false -->
```python
assert type(comp_str('989898989')) == list
```

<!-- grade_id: comp-str-easy, points: 0.5, grade: true, solution: false -->
```python
assert comp_str('1222311') == [(1, 1), (3, 2), (1, 3), (2, 1)]
```

<!-- grade_id: comp-str-hard, points: 0.5, grade: true, solution: false -->
```python
assert comp_str('2') == [(1, 2)]
```

<!-- grade_id: comp-str-empty, points: 0.5, grade: true, solution: false -->
```python
assert comp_str('030303') == [(1, 0), (1, 3), (1, 0), (1, 3), (1, 0), (1, 3)]
```

2. <b>(2 points)</b> __The Mirror of Galadriel:__ With water from the stream Galadriel filled the basin to the brim, and breathed on it, and when the water was still again she spoke. 'Here is the Mirror of Galadriel,' she said. 'I have brought you here so that you may look in it, if you will. For this is what your folk would call magic, I believe; though I do not understand clearly what they mean; and they seem also to use the same word of the deceits of the Enemy. But this, if you will, is the magic of Galadriel. Did you not say that you wished to see Elf-magic?' - Galadriel to Frodo and Sam, describing her Mirror. 

 We call a string S magical if every substring of S appears in Galadriel's Mirror (under lateral inversion). In other words, a magical string is a string where every substring has its reverse in the string. Write a function `magical`, that takes a string `s` as its argument, determine if it is magical or not, and return the result: "Yes" is it is magical, or "No" if it is not magical. For example, `magical('aba')` is 'Yes'. Reason: the list of substrings are : `a, b, ab, ba, aba`. The reverse of each of these strings is present as a substring of S too. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking.

<!-- grade_id: cell-513c33ab888a566f, points: 0, grade: false, solution: true -->
```python
def magical(s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: magical-empty, points: 0.5, grade: true, solution: false -->
```python
assert type(magical("random string")) == str
```

<!-- grade_id: magical-standard, points: 1, grade: true, solution: false -->
```python
assert magical('aba') == 'Yes'
```

<!-- grade_id: magical-hard, points: 0.5, grade: true, solution: false -->
```python
assert magical('magical') == 'No'
```

3. <b>(3 points)</b> Everybody hates Raymond. He’s the largest stamp collector on planet earth and because of that he always makes fun of all the others at the stamp collector parties. Fortunately everybody loves Lucy, and she has a plan. She secretly asks her friends whether they could lend her some stamps, so that she can embarrass Raymond by showing an even larger collection than his. Raymond is so sure about his superiority that he always tells how many stamps he’ll show.And since Lucy knows how many she owns, she knows how many more she needs. She also knows how many friends would lend her some stamps and how many each would lend. But she’s like to borrow from as few friends as possible and if she needs too many then she’d rather not do it at all. Can you tell her the minimum number of friends she needs to borrow from?

 Write a function called `stamps`, which contains the following arguments: how many stamps Lucy needs to borrow and a list containing the number of stamps each of her friends is offering. This function should return the minimum number of friends Lucy needs to borrow stamps from as an integer. If it’s impossible even if she borrows everything from everybody, return "Impossible!" instead. __Note:__ You don't need to worry about error checking for this question, but you are welcome to perform some appropriate error checking.

<!-- grade_id: cell-84a1624b777aea7c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: stampls-long, points: 1, grade: true, solution: false -->
```python
assert type(stamps(4, [1, 1, 1])) == str
assert type(stamps(4, [5])) == int
```

<!-- grade_id: stample-basic, points: 1, grade: true, solution: false -->
```python
assert stamps(100, [13, 17, 42, 9, 23, 57]) == 3
```

<!-- grade_id: stamps-ones, points: 1, grade: true, solution: false -->
```python
assert stamps(4, [1, 1, 1, 1, 1]) == 4
```

4. <b>(2 points) The Bakersville Bake-off!</b> 

The houses in the town of Bakersville are arranged on a grid with R rows (numbered 1 through R) and C columns (numbered 1 through C). The house located on row r and on column c is denoted by (r, c). We say two houses are **neighbors** if they have a common side. 

During the annual Bakersville Bake-off, each household bakes a certain number of cakes to share with their neighbors.
The town comptroller keeps track in an $R \times C$ matrix $A$, where $A_{r, c}$ is the number of cakes which household $(r, c)$ has baked.

The stability of the social order in Bakersville depends on each house baking enough cakes so they can give at least one cake to each neighbor and submit one cake to be judged. Thus, we say a house is **stable** if the number of cakes $A_{r, c}$ is strictly greater than its number of neigbors. We say Bakersville is **stable** if every house is stable. 

The comptroller has asked you to define a function `is_stable` that takes an $R \times C$ numpy array (i.e., an ndarray) as its only argument and returns a Python Boolean that is True if and only if the argument is a stable array and False otherwise. 

<!-- grade_id: cell-4e8532c0e87d3f45, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: basic-example, points: 1, grade: true, solution: false -->
```python
array1 = np.array([4, 7, 8,
                   5, 10, 9,
                   3, 5, 7]).reshape(3, 3)
array2 = np.array([3, 2, 4,
                   1, 1, 2,
                   3, 3, 3]).reshape(3, 3)


assert is_stable(array1)
assert not is_stable(array2)
```

<!-- grade_id: 3-by-3, points: 0.5, grade: true, solution: false -->
```python
assert not is_stable(np.array([1., 0., 1.]))
```

<!-- grade_id: 4-by-4, points: 0.5, grade: true, solution: false -->
```python
assert is_stable(np.array([[4, 6, 6, 4], 
                           [4, 7, 7, 7], 
                           [4, 21, 8, 4], 
                           [4, 4, 7, 5]]))
```